# **Save this file as studentid1_studentid2_lab#.ipynb**
(Your student-id is the number shown on your student card.)

E.g. if you work with 3 people, the notebook should be named:
12301230_3434343_1238938934_lab1.ipynb.

**This will be parsed by a regexp, so please double check your filename.**

Before you turn this problem in, please make sure everything runs correctly. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

**Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your names and email adresses below.**



```python
NAME = "Joris Mollinga"
NAME2 = "Ruben Polak"
NAME3 = ""
EMAIL = "joris.mollinga@student.uva.nl"
EMAIL2 = "rubenpolak@gmail.com"
EMAIL3 = ""
```

---

# Lab 1: Linear Regression and Overfitting

### Machine Learning 1, September 2017

Notes on implementation:

* You should write your code and answers in this IPython Notebook: http://ipython.org/notebook.html. If you have problems, please contact your teaching assistant.
* Please write your answers right below the questions.
* Among the first lines of your notebook should be "%pylab inline". This imports all required modules, and your plots will appear inline.
* Refer to last week's lab notes, i.e. http://docs.scipy.org/doc/, if you are unsure about what function to use. There are different correct ways to implement each problem!
* For this lab, your regression solutions should be in closed form, i.e., should not perform iterative gradient-based optimization but find the exact optimum directly.
* use the provided test boxes to check if your answers are correct

```python
%pylab inline
plt.rcParams["figure.figsize"] = [20,10]
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

$\newcommand{\bPhi}{\mathbf{\Phi}}$
$\newcommand{\bx}{\mathbf{x}}$
$\newcommand{\bw}{\mathbf{w}}$
$\newcommand{\bt}{\mathbf{t}}$
$\newcommand{\by}{\mathbf{y}}$
$\newcommand{\bm}{\mathbf{m}}$
$\newcommand{\bS}{\mathbf{S}}$
$\newcommand{\bI}{\mathbf{I}}$

## Part 1: Polynomial Regression

### 1.1. Generate periodic data (5 points)
Write a method `gen_cosine(N)` that generates toy data like in fig 1.2 of Bishop's book. The method should have a parameter $N$, and should return $N$-dimensional vectors $\bx$ and $\bt$, where $\bx$ contains evenly spaced values from 0 to (including) 2$\pi$, and the elements $t_i$ of $\bt$ are distributed according to:

$$t_i \sim \mathcal{N}(\mu_i, \sigma^2)$$

where $x_i$ is the $i$-th elements of $\bf{x}$, the mean $\mu_i = \cos(x_i)$ and the standard deviation $\sigma = 0.2$.


```python
import numpy as np
```

<!-- grade_id: cell-1c8c68d862f80f7e, points: 0, grade: false, solution: true -->
```python
def gen_cosine(n):
    import numpy as np
    # YOUR CODE HERE
    x = np.linspace(0,2*np.pi,n)
  
    t = np.random.normal(np.cos(x),0.2)
    
    # raise NotImplementedError()
    
    #plt.scatter(x,t)
    #plt.plot(x,cos(x))
    # print (len(x) == N)
    return x, t


```

<!-- grade_id: cell-afc3c8025a62af85, points: 5, grade: true, solution: false -->
```python
### Test your function
np.random.seed(5)
N = 10
x, t = gen_cosine(N)

assert x.shape == (N,), "the shape of x is incorrect"
assert t.shape == (N,), "the shape of t is incorrect"

```

### 1.2 Polynomial regression (10 points)

Write a method `fit_polynomial(x, t, M)` that finds the maximum-likelihood solution of an _unregularized_ $M$-th order polynomial for some dataset `x`. The error function to minimize w.r.t. $\bw$ is:

$E(\bw) = \frac{1}{2} (\bPhi\bw - \bt)^T(\bPhi\bw - \bt)$

where $\bPhi$ is the _feature matrix_ (or _design matrix_) as explained in Bishop's book at section 3.1.1, $\bt$ is the vector of target values. Your method should return a vector $\bw$ with the maximum-likelihood parameter estimates, as well as the _feature matrix_ $\bPhi$.

<!-- grade_id: cell-27b0d6f1bcdeb97b, points: 0, grade: false, solution: true -->
```python
def convertToArray(x):
    length = len(x)
    
    xTemp = np.zeros([length,1])
    
    for i in range(0,len(x)):
        xTemp[i] = x[i]

    return xTemp


def designmatrix(x, M): # it is highly recommended to write a helper function that computes Phi
    # YOUR CODE HERE
    N = len(x)
    if len(x) != N:
        raise NotImplementedError()
        
    Phi = np.zeros([N,M+1])
    
    for row in range (N):
        for col in range(M+1):
            Phi[row,col] = x[row]**col
     
    return Phi


def fit_polynomial(x, t, M):
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    Phi = designmatrix(x,M)
    
    t = convertToArray(t)
    
    w_ml = np.dot(np.dot(np.linalg.inv(np.dot(np.transpose(Phi),Phi)),np.transpose(Phi)),t)
    
    w_ml2 = w_ml.ravel()
    
    return w_ml2, Phi

```

<!-- grade_id: cell-7e0f87e75d7c82aa, points: 10, grade: true, solution: false -->
```python
### Test your function
N = 10
x = np.square((np.linspace(-1, 1, N)))
t = 0.5*x + 1.5
m = 2
w, Phi = fit_polynomial(x,t,m)

assert w.shape == (m+1,), "The shape of w is incorrect"
assert Phi.shape == (N, m+1), "The shape of Phi is incorrect"

```

### 1.3 Plot (5 points)
Sample a dataset with $N=10$, and fit four polynomials with $M \in (0, 2, 4, 8)$.
For each value of $M$, plot the prediction function, along with the data and the original cosine function. The resulting figure should look similar to fig 1.4 of the Bishop's book. Note that you can use matplotlib's `plt.pyplot(.)` functionality for creating grids of figures.

<!-- grade_id: cell-2774a098ae80cb7e, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
def calculateFunctionValues(x,M,w):
    if M == 0:
        return w[0]
    
    val = 0
    for j in range(1,M+1):
        val = val + w[j]*(x**j)
    val = val + w[0]
    
    return val

def createFunctions(x,t,N):
    
    x_smooth = np.linspace(0,2*np.pi,N*100)
    
    t1 = np.zeros(N*100)
    t2 = np.zeros(N*100)
    t3 = np.zeros(N*100)
    t4 = np.zeros(N*100)
    
    M1,M2,M3,M4 = 0,2,4,8
    
    w1,Phi1 = fit_polynomial(x,t,M1)
    w2,Phi2 = fit_polynomial(x,t,M2)
    w3,Phi3 = fit_polynomial(x,t,M3)
    w4,Phi4 = fit_polynomial(x,t,M4)
    
    for i in range(0,N*100):
        t1[i] = calculateFunctionValues(x_smooth[i],M1,w1)
        t2[i] = calculateFunctionValues(x_smooth[i],M2,w2)
        t3[i] = calculateFunctionValues(x_smooth[i],M3,w3)
        t4[i] = calculateFunctionValues(x_smooth[i],M4,w4)
    return t1,t2,t3,t4,x_smooth

N = 10
x,t = gen_cosine(N)

t1,t2,t3,t4,x_smooth = createFunctions(x,t,N)
t_cos = np.cos(x_smooth)

def plotFunctions(t1,t2,t3,t4,x_smooth,x,t):
    f, axarr = plt.subplots(2, 2)
    axarr[0, 0].plot(x_smooth,t_cos)
    axarr[0, 0].scatter(x, t)
    axarr[0, 0].plot(x_smooth, t1)
    axarr[0, 0].text(3, 0.2 , "M=0",fontsize=18)
    axarr[0, 1].plot(x_smooth,t_cos)
    axarr[0, 1].scatter(x, t)
    axarr[0, 1].plot(x_smooth, t2)
    axarr[0, 1].text(3, 0.2 , "M=2",fontsize=18)
    axarr[1, 0].plot(x_smooth,t_cos)
    axarr[1, 0].scatter(x, t)
    axarr[1, 0].plot(x_smooth, t3)
    axarr[1, 0].text(3, 0.2 , "M=4",fontsize=18)
    axarr[1, 1].plot(x_smooth,t_cos)
    axarr[1, 1].scatter(x, t)
    axarr[1, 1].plot(x_smooth, t4)
    axarr[1, 1].text(3, 0.2 , "M=8",fontsize=18)
    return


### UNCOMMENT THE LINE BELOW TO PLOT THE FUNCTIONS
# plotFunctions(t1,t2,t3,t4,x_smooth,x,t)
```

### 1.4 Regularized linear regression (10 points)

Write a method `fit_polynomial_reg(x, t, M, lamb)` that fits a _regularized_ $M$-th order polynomial to the periodic data, as discussed in the lectures, where `lamb` is the regularization term _lambda_. (Note that 'lambda' cannot be used as a variable name in Python since it has a special meaning). The error function to minimize w.r.t. $\bw$:

$E(\bw) = \frac{1}{2} (\bPhi\bw - \bt)^T(\bPhi\bw - \bt) + \frac{\lambda}{2} \mathbf{w}^T \mathbf{w}$

For background, see section 3.1.4 of Bishop's book.

The function should return $\bw$ and $\bPhi$.

<!-- grade_id: cell-9fe81fd5537daff9, points: 0, grade: false, solution: true -->
```python
def fit_polynomial_reg(x, t, m, lamb):
    Phi = designmatrix(x,m)
    t = convertToArray(t)
    w_ml = np.dot(np.linalg.inv(np.dot(lamb,np.identity(m+1)) + np.dot(np.transpose(Phi),Phi)),(np.dot(np.transpose(Phi),t)))
    #raise NotImplementedError()
    w_ml = w_ml.ravel()
    return w_ml, Phi

```

<!-- grade_id: cell-a2e561cbc5a4140b, points: 10, grade: true, solution: false -->
```python
### Test your function
N = 10
x = np.square((np.linspace(-1, 1, N)))
t = 0.5*x + 1.5
m = 2
lamb = 0.1
w, Phi = fit_polynomial_reg(x,t,m, lamb)

assert w.shape == (m+1,), "The shape of w is incorrect"
assert Phi.shape == (N, m+1), "The shape of w is incorrect" 

```

```python
N = 10
x,t = gen_cosine(N)
M = 8
lamb = 0.1

w_ml,Phi = fit_polynomial_reg(x,t,M,lamb)
```






### 1.5 Model selection by cross-validation (15 points)
Use cross-validation to find a good choice of $M$ and $\lambda$, given a dataset of $N=10$ datapoints generated with `gen_cosine(20)`. You should write a function that tries (loops over) a reasonable range of choices of $M$ and $\lambda$, and returns the choice with the best cross-validation error. In this case you use $K=5$ folds.

You can let $M \in (0, 1, ..., 10)$, and let $\lambda \in (e^{-10}, e^{-9}, ..., e^{0})$.

a) (5 points) First of all, write a method `pred_error(x_train, x_valid, t_train, t_valid, M, lamb)` that compares the prediction of your method `fit_polynomial_reg` for a given set of parameters $M$ and $\lambda$ to `t_valid`. It should return the prediction error for a single fold.

<!-- grade_id: cell-d631a845dec603be, points: 0, grade: false, solution: true -->
```python
def pred_error(x_train, x_valid, t_train, t_valid, M, reg):
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    
    
    return pred_err

```

<!-- grade_id: cell-ba7261e2eae040fd, points: 5, grade: true, solution: false -->
```python
### Test your function
N = 10
x = np.linspace(-1, 1, N)
t = 0.5*np.square(x) + 1.5
M = 2
reg = 0.1
pred_err = pred_error(x[:-2], x[-2:], t[:-2], t[-2:], M, reg)

assert pred_err < 0.01, "pred_err is too big"

```

Output:
```
<ipython-input-12-a22a68e3f8f3> in <module>()
      5 M = 2
      6 reg = 0.1
----> 7 pred_err = pred_error(x[:-2], x[-2:], t[:-2], t[-2:], M, reg)
      8 
      9 assert pred_err < 0.01, "pred_err is too big"

<ipython-input-11-f518f8c85cfb> in pred_error(x_train, x_valid, t_train, t_valid, M, reg)
      1 def pred_error(x_train, x_valid, t_train, t_valid, M, reg):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()
      4 
      5     return pred_err

NotImplementedError:
```

b) (10 points) Now write a method find_best_m_and_lamb(x, t) that finds the best values for $M$ and $\lambda$. The method should return the best $M$ and $\lambda$. To get you started, here is a method you can use to generate indices of cross-validation folds.

```python
def kfold_indices(N, k):
    all_indices = np.arange(N,dtype=int)
    np.random.shuffle(all_indices)
    idx = [int(i) for i in np.floor(np.linspace(0,N,k+1))]
    train_folds = []
    valid_folds = []
    for fold in range(k):
        valid_indices = all_indices[idx[fold]:idx[fold+1]]
        valid_folds.append(valid_indices)
        train_folds.append(np.setdiff1d(all_indices, valid_indices))
    return train_folds, valid_folds

```

<!-- grade_id: cell-0553f08188fbfcd6, points: 0, grade: false, solution: true -->
```python
def find_best_m_and_lamb(x, t):
    # YOUR CODE HERE
    # check
    raise NotImplementedError()
    
    return M_best, lamb_best

```

<!-- grade_id: cell-523aa38e51c8913c, points: 10, grade: true, solution: false -->
```python
### If you want you can write your own test here
```

### 1.7 Plot best cross-validated fit (5 points)

For some dataset with $N = 10$, plot the model with the optimal $M$ and $\lambda$ according to the cross-validation error, using the method you just wrote. In addition, the plot should show the dataset itself and the function that we try to approximate. Let the plot make clear which $M$ and $\lambda$ were found.

<!-- grade_id: cell-f59cac6f24ce02fc, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part 2: Bayesian Linear (Polynomial) Regression

### 2.1 Cosine 2 (5 points)

Write a function `gen_cosine2(N)` that behaves identically to `gen_cosine(N)` except that the generated values $x_i$ are not linearly spaced, but drawn from a uniform distribution between $0$ and $2 \pi$.

<!-- grade_id: cell-36addc88e1b3fe1d, points: 0, grade: false, solution: true -->
```python
def gen_cosine2(n):
    
    x = np.random.uniform(0, 2*np.pi, N)
    x.sort()
    t= np.random.normal(np.cos(x),0.2)
    return x, t

```

<!-- grade_id: cell-57a847b66fb591af, points: 5, grade: true, solution: false -->
```python
### Test your function
np.random.seed(5)
N = 10
x, t = gen_cosine2(N)

assert x.shape == (N,), "the shape of x is incorrect"
assert t.shape == (N,), "the shape of t is incorrect"

```

### 2.2 Compute Posterior (15 points)

You're going to implement a Bayesian linear regression model, and fit it to the periodic data. Your regression model has a zero-mean isotropic Gaussian prior over the parameters, governed by a single (scalar) precision parameter $\alpha$, i.e.:

$$p(\bw \;|\; \alpha) = \mathcal{N}(\bw \;|\; 0, \alpha^{-1} \bI)$$

The covariance and mean of the posterior are given by:

$$\bS_N= \left( \alpha \bI + \beta \bPhi^T \bPhi \right)^{-1} $$
$$\bm_N = \beta\; \bS_N \bPhi^T \bt$$

where $\alpha$ is the precision of the predictive distribution, and $\beta$ is the noise precision. 
See MLPR chapter 3.3 for background.

Write a method `fit_polynomial_bayes(x, t, M, alpha, beta)` that returns the mean $\bm_N$ and covariance $\bS_N$ of the posterior for a $M$-th order polynomial. In addition it should return the design matrix $\bPhi$. The arguments `x`, `t` and `M` have the same meaning as in question 1.2.

<!-- grade_id: cell-24d68ce462db8f40, points: 0, grade: false, solution: true -->
```python
def designmatrix(x, M):
    
    Phi = np.zeros((N,M+1))
    for row in range (N):
        for col in range(M+1):
            Phi[row,col] = x[row]**col 
    
    return Phi, np.transpose(Phi)

def fit_polynomial_bayes(x, t, M, alpha, beta):

    Phi, Phi_T = designmatrix(x, M)
    S = np.linalg.inv(np.dot(alpha, np.identity(M+1)) + np.dot(beta,np.dot(Phi_T,Phi)))
    m = np.dot(beta, np.dot(S, np.dot(Phi_T, t)))

    if not m.any() or not S.any() or not Phi.any():
        raise NotImplementedError()
    
    return m, S, Phi
```

<!-- grade_id: cell-5cf266fca46cd4e9, points: 15, grade: true, solution: false -->
```python
### Test your function
N = 10
x = np.linspace(-1, 1, N)
t = 0.5*np.square(x) + 1.5
M = 2
alpha = 0.5
beta = 25
m, S, Phi = fit_polynomial_bayes(x, t, M, alpha, beta)

assert m.shape == (M+1,), "the shape of m is incorrect" 
assert S.shape == (M+1, M+1), "the shape of S is incorrect"
assert Phi.shape == (N, M+1), "the shape of Phi is incorrect"

```

### 2.3 Prediction (10 points)

The predictive distribution of Bayesian linear regression is:

$$ p(t \;|\; \bx, \bt, \alpha, \beta) = \mathcal{N}(t \;|\; \bm_N^T \phi(\bx), \sigma_N^2(\bx))$$

$$ \sigma_N^2 = \frac{1}{\beta} + \phi(\bx)^T \bS_N \phi(\bx) $$

where $\phi(\bx)$ are the computed features for a new datapoint $\bx$, and $t$ is the predicted variable for datapoint $\bx$. 

Write a function that `predict_polynomial_bayes(x, m, S, beta)` that returns the predictive mean, variance and design matrix $\bPhi$ given a new datapoint `x`, posterior mean `m`, posterior variance `S` and a choice of model variance `beta`.

<!-- grade_id: cell-865ac2a455500d4b, points: 0, grade: false, solution: true -->
```python
def predict_polynomial_bayes(x, m, S, beta):
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return mean, sigma, Phi

```

<!-- grade_id: cell-ecb92bcbe0131eb9, points: 10, grade: true, solution: false -->
```python
### Test your function
np.random.seed(5)
N = 10
x = np.linspace(-1, 1, N)
m = np.empty(3)
S = np.empty((3, 3))
beta = 25
mean, sigma, Phi = predict_polynomial_bayes(x, m, S, beta)

assert mean.shape == (N,), "the shape of mean is incorrect"
assert sigma.shape == (N,), "the shape of sigma is incorrect"
assert Phi.shape == (N, m.shape[0]), "the shape of Phi is incorrect"

```

### 2.4 Plot predictive distribution (10 points)

a) (5 points) Generate 10 datapoints with `gen_cosine2(10)`. Compute the posterior mean and covariance for a Bayesian polynomial regression model with $M=4$, $\alpha=\frac{1}{2}$ and $\beta=\frac{1}{0.2^2}$.
Plot the Bayesian predictive distribution, where you plot (for $x$ between 0 and $2 \pi$) $t$'s predictive mean and a 1-sigma predictive variance using `plt.fill_between(..., alpha=0.1)` (the alpha argument induces transparency).

Include the datapoints in your plot.


<!-- grade_id: cell-81339ee6f9873831, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

b) (5 points) For a second plot, draw 100 samples from the parameters' posterior distribution. Each of these samples is a certain choice of parameters for 4-th order polynomial regression. 
Display each of these 100 polynomials.


<!-- grade_id: cell-68cf4e7a773b42b9, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### 2.5 Additional questions (10 points)

a) (5 points) Why is $\beta=\frac{1}{0.2^2}$ the best choice of $\beta$ in section 2.4?



<!-- grade_id: cell-ac9287d87d554547, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

b) (5 points) What problems do we face when it comes to choosing basis functions in linear models?

<!-- grade_id: cell-a3e38ba4c988b6a9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE